# PART A: Selecting the right ML Model
© 2023 Zaka AI, Inc. All Rights Reserved.

---

Please report your choices for part A in the code cells below.

**NOTE:**
* For the algo_type, you answer by 'supervised' or 'unsupervised'.
* For the task_type, you answer by 'regression', 'classification', 'clustering', or 'association analysis'
* For the suggested algorithms you pick one or more from the following list AS THEY APPEAR. ['linear regression', 'multilinear regression', 'decision tree Classifier', 'svm', 'logistic regression', 'knn', 'random forest classifier', 'naive bayes', 'k-means clustering', 'apriori']

Even if you select one algorithm, make sure to write it inside a list (as written below)

Sample answer:

algo_type_x = 'supervised' </br>
task_type_x = 'classification' </br>
algos_x = ['naive bayes']

```python
# For the Telecom Users dataset:
algo_type_1 = 'supervised'
task_type_1 = 'regression'
algos_1 = ['multilinear regression']
```

```python
#For the Mobile Prices dataset:
algo_type_2 = 'supervised'
task_type_2 = 'classification'
algos_2 = ['svm', 'random forest classifier', 'naive bayes', 'knn', 'decision tree classifier']
```

```python
#For the Iris dataset:
algo_type_3 = 'supervised'
task_type_3 = 'classification'
algos_3 = ['naive bayes', 'svm', 'random forest classifier', 'knn', 'decision tree classifier']
```

```python
#For the student performance dataset:
#I believe this is a trick question, check xls file.
algo_type_4 = 'unsupervised'
task_type_4 = 'clustering'
algos_4 = ['k-means clustering']
```

```python
#For the stroke dataset, the algorithm type is:
algo_type_5 = 'supervised'
task_type_5 = 'classification'
algos_5 = ['logistic regression', 'naive bayes', 'random forest classifier', 'decision tree classifier', 'svm', 'knn']
```


# PART B: Gaussian Naive Bayes Classifier

© 2023 Zaka AI, Inc. All Rights Reserved.

---



## **Case Study:** Iris Dataset

**Objective:** The objective of this challenge is to make you know about Naive Bayes applied on Numerical Values.

**DataSet Columns:**<br>
*	 Petal Height
*  Petal Width
*  Sepal Height
*  Sepal Width
*  Target: The kind of the Iris flower (Virginica, Setosa, Versicolor)

# Importing Libraries

```python
import numpy as np
import pandas as pd
import math
```

```python
import warnings
warnings.filterwarnings('ignore')
```

# Loading the Dataset

Load the dataset, and make sure your columns are named as follow: 'Sepal Height', 'Sepal Width', 'Petal Height', 'Petal Width', and 'Target'

```python
# clone git repo
!git clone https://github.com/royaad/AIC
# change working directory
%cd AIC/Week_03/data
```

Output:
```
Cloning into 'AIC'...
remote: Enumerating objects: 25, done.[K
remote: Counting objects: 100% (25/25), done.[K
remote: Compressing objects: 100% (16/16), done.[K
remote: Total 25 (delta 1), reused 0 (delta 0), pack-reused 0[K
Unpacking objects: 100% (25/25), 12.73 KiB | 2.12 MiB/s, done.
/content/AIC/Week_03/data
```

```python
filename = 'iris.csv'
df = pd.read_csv(filename)
#replacing dots with spaces in columns
df.columns = df.columns.str.replace('.',' ')
#replacing variety with Target & length with Height by chaining
df.columns = df.columns.str.replace('variety','Target').str.replace('length','Height')
#capitalizing first letters
df.columns = df.columns.str.title()
#showing first 5 rows
df.head()
```

Output:
```
Sepal Height  Sepal Width  Petal Height  Petal Width  Target
0           5.1          3.5           1.4          0.2  Setosa
1           4.9          3.0           1.4          0.2  Setosa
2           4.7          3.2           1.3          0.2  Setosa
3           4.6          3.1           1.5          0.2  Setosa
4           5.0          3.6           1.4          0.2  Setosa
```

## Data Preprocessing

We will Change the string values of the Target column to numerical.

```python
col = 'Target'
target_dict = {ele:count for count, ele in enumerate(df[col].unique())}
df[col] = df[col].apply(lambda x: target_dict[x])
print(target_dict)
df.head(10)
```

Output:
```
{'Setosa': 0, 'Versicolor': 1, 'Virginica': 2}
   Sepal Height  Sepal Width  Petal Height  Petal Width  Target
0           5.1          3.5           1.4          0.2       0
1           4.9          3.0           1.4          0.2       0
2           4.7          3.2           1.3          0.2       0
3           4.6          3.1           1.5          0.2       0
4           5.0          3.6           1.4          0.2       0
5           5.4          3.9           1.7          0.4       0
6           4.6          3.4           1.4          0.3       0
7           5.0          3.4           1.5          0.2       0
8           4.4          2.9           1.4          0.2       0
9           4.9          3.1           1.5          0.1       0
```

Make sure we have no null values

```python
#Test Your Zaka
print ("Missing values per column:\n")
print (df.isnull().sum(axis = 0))
```

Output:
```
Missing values per column:

Sepal Height    0
Sepal Width     0
Petal Height    0
Petal Width     0
Target          0
dtype: int64
```

# Naive Bayes

## Finding different Classes

First, let's try to find how many classes we have in our dataset, although it should always appear in the description of your dataset.

```python
targets = len(df[col].unique())
print(f'The # of classes is {targets}')
```

Output:
```
The # of classes is 3
```

So, we have 3 classes of flowers.

Remember the basic formula that we used for Naive Bayes. <br>
<img src="https://equatio-api.texthelp.com/svg/%5C%20P(%5Ctextcolor%7B%232B7FBB%7D%7BClass%7D%7C%5Ctextcolor%7B%23E94D40%7D%7BFeatures%7D)%3D%5Cfrac%7BP(%5Ctextcolor%7B%23E94D40%7D%7BFeatures%7D%7C%5Ctextcolor%7B%232B7FBB%7D%7BClass%7D)%5Ccdot%20P%5Cleft(%5Ctextcolor%7B%232B7FBB%7D%7BClass%7D%5Cright)%7D%7BP(%5Ctextcolor%7B%23E94D40%7D%7BFeatures%7D)%7D" alt="P of open paren C l a. s s divides F of e a. t u r e s close paren equals the fraction with numerator P of open paren F of e a. t u r e s divides C l a. s s close paren times P of open paren C l a. s s close paren and denominator P of F of e a. t u r e s">

Since we have 3 classes, and 4 features, we need to calculate the following probabilities.<br>
<img src="https://equatio-api.texthelp.com/svg/P(%5Ctextcolor%7B%232B7FBB%7D%7BClass_0%7D%7C%5Ctextcolor%7B%23E94D40%7D%7BF1%2CF2%2CF3%2CF4%7D)" alt="P of open paren C l a. s s sub 0 divides F of 1 comma F of 2 comma F of 3 comma F of 4 close paren"> <br>
<img src="https://equatio-api.texthelp.com/svg/P(%5Ctextcolor%7B%232B7FBB%7D%7BClass_1%7D%7C%5Ctextcolor%7B%23E94D40%7D%7BF1%2CF2%2CF3%2CF4%7D)" alt="P of open paren C l a. s s sub 1 divides F of 1 comma F of 2 comma F of 3 comma F of 4 close paren"> <br>
<img src="https://equatio-api.texthelp.com/svg/P(%5Ctextcolor%7B%232B7FBB%7D%7BClass_2%7D%7C%5Ctextcolor%7B%23E94D40%7D%7BF1%2CF2%2CF3%2CF4%7D)" alt="P of open paren C l a. s s sub 2 divides F of 1 comma F of 2 comma F of 3 comma F of 4 close paren">


So in reality we need to calculate the following:

<img src="https://equatio-api.texthelp.com/svg/P_0%3DP(%5Ctextcolor%7B%232B7FBB%7D%7B%5Ctextcolor%7B%23E94D40%7D%7BF_1%7D%7D%7C%5Ctextcolor%7B%23E94D40%7D%7B%5Ctextcolor%7B%232B7FBB%7D%7BClass_0%7D%7D)P(%5Ctextcolor%7B%232B7FBB%7D%7B%5Ctextcolor%7B%23E94D40%7D%7BF_2%7D%7D%7C%5Ctextcolor%7B%23E94D40%7D%7B%5Ctextcolor%7B%232B7FBB%7D%7BClass_0%7D%7D)P(%5Ctextcolor%7B%232B7FBB%7D%7B%5Ctextcolor%7B%23E94D40%7D%7BF_3%7D%7D%7C%5Ctextcolor%7B%23E94D40%7D%7B%5Ctextcolor%7B%232B7FBB%7D%7BClass_0%7D%7D)P(%5Ctextcolor%7B%232B7FBB%7D%7B%5Ctextcolor%7B%23E94D40%7D%7BF_4%7D%7D%7C%5Ctextcolor%7B%23E94D40%7D%7B%5Ctextcolor%7B%232B7FBB%7D%7BClass_0%7D%7D)" alt="P sub 0 equals P of open paren F sub 1 divides C l a. s s sub 0 close paren P of open paren F sub 2 divides C l a. s s sub 0 close paren P of open paren F sub 3 divides C l a. s s sub 0 close paren P of open paren F sub 4 divides C l a. s s sub 0 close paren"><img src="https://equatio-api.texthelp.com/svg/P%5Cleft(%5Ctextcolor%7B%232B7FBB%7D%7BClass_0%7D%5Cright)" alt="P of open paren C l a. s s sub 0 close paren"><br><img src="https://equatio-api.texthelp.com/svg/P_1%3DP(%5Ctextcolor%7B%232B7FBB%7D%7B%5Ctextcolor%7B%23E94D40%7D%7BF_1%7D%7D%7C%5Ctextcolor%7B%23E94D40%7D%7B%5Ctextcolor%7B%232B7FBB%7D%7BClass_1%7D%7D)P(%5Ctextcolor%7B%232B7FBB%7D%7B%5Ctextcolor%7B%23E94D40%7D%7BF_2%7D%7D%7C%5Ctextcolor%7B%23E94D40%7D%7B%5Ctextcolor%7B%232B7FBB%7D%7BClass_1%7D%7D)P(%5Ctextcolor%7B%232B7FBB%7D%7B%5Ctextcolor%7B%23E94D40%7D%7BF_3%7D%7D%7C%5Ctextcolor%7B%23E94D40%7D%7B%5Ctextcolor%7B%232B7FBB%7D%7BClass_1%7D%7D)P(%5Ctextcolor%7B%232B7FBB%7D%7B%5Ctextcolor%7B%23E94D40%7D%7BF_4%7D%7D%7C%5Ctextcolor%7B%23E94D40%7D%7B%5Ctextcolor%7B%232B7FBB%7D%7BClass_1%7D%7D)" alt="P sub 1 equals P of open paren F sub 1 divides C l a. s s sub 1 close paren P of open paren F sub 2 divides C l a. s s sub 1 close paren P of open paren F sub 3 divides C l a. s s sub 1 close paren P of open paren F sub 4 divides C l a. s s sub 1 close paren"><img src="https://equatio-api.texthelp.com/svg/P%5Cleft(%5Ctextcolor%7B%232B7FBB%7D%7BClass_1%7D%5Cright)" alt="P of open paren C l a. s s sub 1 close paren"><br>
<img src="https://equatio-api.texthelp.com/svg/P_2%3DP(%5Ctextcolor%7B%232B7FBB%7D%7B%5Ctextcolor%7B%23E94D40%7D%7BF_1%7D%7D%7C%5Ctextcolor%7B%23E94D40%7D%7B%5Ctextcolor%7B%232B7FBB%7D%7BClass_2%7D%7D)P(%5Ctextcolor%7B%232B7FBB%7D%7B%5Ctextcolor%7B%23E94D40%7D%7BF_2%7D%7D%7C%5Ctextcolor%7B%23E94D40%7D%7B%5Ctextcolor%7B%232B7FBB%7D%7BClass_2%7D%7D)P(%5Ctextcolor%7B%232B7FBB%7D%7B%5Ctextcolor%7B%23E94D40%7D%7BF_3%7D%7D%7C%5Ctextcolor%7B%23E94D40%7D%7B%5Ctextcolor%7B%232B7FBB%7D%7BClass_2%7D%7D)P(%5Ctextcolor%7B%232B7FBB%7D%7B%5Ctextcolor%7B%23E94D40%7D%7BF_4%7D%7D%7C%5Ctextcolor%7B%23E94D40%7D%7B%5Ctextcolor%7B%232B7FBB%7D%7BClass_2%7D%7D)P%5Cleft(%5Ctextcolor%7B%232B7FBB%7D%7BClass_2%7D%5Cright)" alt="P sub 2 equals P of open paren F sub 1 divides C l a. s s sub 2 close paren P of open paren F sub 2 divides C l a. s s sub 2 close paren P of open paren F sub 3 divides C l a. s s sub 2 close paren P of open paren F sub 4 divides C l a. s s sub 2 close paren P of open paren C l a. s s sub 2 close paren">



We see which one is the greatest, and based on that we assign the class.

Those probabilities will be approximated using a distribution.
In this example, we will use the Gaussien Distribution.

## Gaussian Probability Density Function

We recall that the Gaussian Probability density function is given by:
<br>
<img src="https://equatio-api.texthelp.com/svg/f%5Cleft(x%5Cright)%3D%5Cfrac%7B1%7D%7B%5Csqrt%7B2%5Cpi%7D%5Ctextcolor%7B%238D44AD%7D%7B%5Csigma%7D%7D%5Cexp%5Cleft%5C%7B-%5Cfrac%7B%5Cleft(x-%5Ctextcolor%7B%233697DC%7D%7Bmean%7D%5Cright)%5E2%7D%7B2%5Ctextcolor%7B%238D44AD%7D%7B%5Csigma%7D%5E2%7D%5Cright%5C%7D" alt="f of x equals 1 over the square root of 2 pi sigma the exp of open brace negative the fraction with numerator open paren x minus m e a. n close paren squared and denominator 2 sigma squared close brace">

```python
def probability (x, mean=0, sigma=1):

  A = (1/math.sqrt(2*math.pi*sigma))
  prob = A*np.exp(-(x - mean)**2 / (2*sigma**2))

  return prob
```

## Naive Bayes Implementation

Let's implement Naive Bayes function and then test it with a prediction

```python
def naive_bayes (df, features, target_name):
  # A function that receives a dataframe, a feature vector, and a target name, and retruns the most likely class for the target_name given the feature vector and the dataframe.
  # df: the dataframe from which we derive the probabilities.
  # features: a 1 dimensional numpy array that contains the features of 1 sample we want to test
  # target_name: the column name inside the dataframe df that represents the target we want to predict.
  # Refer to the code cell below to see how this function will be used and accordingly code it.


  # target column.
  classes = np.int_(df[:,-1])
  # counting the total number of examples.
  N_classes = len(classes)
  # initiationg some lists.
  p_classes = []
  features_means = np.empty((len(np.unique(classes)),df.shape[1]-1))
  features_means[:] = np.nan
  features_stdev = features_means.copy()

  for i in np.unique(classes):

    # calculating the probability for each class.
    p_class = len(classes[classes == i])/N_classes
    p_classes.append(p_class)
    # calculating the mean and std for each feature for each class.
    temp = df[classes == i,:-1]
    features_means[i,:] = temp.mean(axis=0)
    features_stdev[i,:] = temp.std(axis=0)

  # treating 1 set of features case
  if len(features.shape) == 1:
    # intitiating proba
    p_target = p_classes
    for i in np.unique(classes):
      for k in range(len(features)):
           p_target[i] *= probability(features[k],features_means[i,k],features_stdev[i,k])
    result = np.argmax(p_target)
  # treating multiple stes
  # previous lines of code can be reduced by converting the 1D array into 2D array.
  else:
    # initiating the probabilty
    p_target = np.tile(p_classes,(features.shape[0],1))
    for k in range(features.shape[0]):
      for l in np.unique(classes):
        for m in range(features.shape[1]):
           p_target[k,l] *= probability(features[k,m],features_means[l,m],features_stdev[l,m])
    result = np.argmax(p_target,axis = 1)

  return result
```

```python
features = np.array([4.9, 3.0, 1.4, 0.2])
result = naive_bayes(df.values, features, 'Target')
print("This Iris flower is in the class ",result)
```

Output:
```
This Iris flower is in the class  0
```

See the performance of our NB model

Now here we will splot our data between 2 sets:

*   One from which the Naive Bayes Model will take the probabilities. (The **train** set)
*   one that it hasn't seen before to test on it (The **test** set)
Split your dataset as 80/20 with a random_state of 0

```python
x_train, x_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=0)from sklearn.model_selection import train_test_split

#Test Your Zaka
# store all columns except last one as inputs in X
X = df.iloc[:,0:-1].values
# store the last column as the output (label) in y
y = df.iloc[:,-1].values

# split dataset in a 80/20 split
x_train, x_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=0)
print(x_train.shape, y_train.shape)
print(x_test.shape, y_test.shape)
```

Output:
```
(120, 4) (120,)
(30, 4) (30,)
```

Compute the accuracy of Naive Bayes on the test set.

```python
train_dataFrame = np.c_[x_train,y_train] #Concatenate your x_train and y_train into 1 dataframe because you will pass it to the naive_bayes function you coded.

# predicting results for test set
y_pred = naive_bayes(train_dataFrame, x_test, 'Target')

print(f'predicted ys are \n {np.array(y_pred)}')
print(f'test ys are \n {y_test}')
y_wrong = np.count_nonzero(y_pred - y_test, axis = 0)
accuracy = 100*(len(y_test) - y_wrong)/len(y_test)

print("Accuracy of our Naive Bayes is:", accuracy, "%")
```

Output:
```
predicted ys are 
 [2 1 0 2 0 2 0 1 1 1 2 1 1 1 1 0 1 1 0 0 2 1 0 0 2 0 0 1 1 0]
test ys are 
 [2 1 0 2 0 2 0 1 1 1 2 1 1 1 1 0 1 1 0 0 2 1 0 0 2 0 0 1 1 0]
Accuracy of our Naive Bayes is: 100.0 %
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. **Download your notebook as .ipynb file and put it in a zip folder along with the PART A explanation, then upload that folder to the platform**

Tests for q0 (points: 50)
```
>>> 'multilinear regression' in algos_1 and (not bool(set(algos_1) & set(['decision tree Classifier', 'svm', 'logistic regression', 'knn', 'random forest classifier', 'naive bayes', 'k-means clustering', 'apriori', 'linear regression'])))
True
```
```
>>> bool(set(algos_2) & set(['decision tree Classifier', 'svm', 'logistic regression', 'knn', 'random forest classifier', 'naive bayes'])) and (not bool(set(algos_2) & set(['k-means clustering', 'apriori', 'linear regression','multilinear regression'])))
True
```
```
>>> bool(set(algos_3) & set(['decision tree Classifier', 'svm', 'logistic regression', 'knn', 'random forest classifier', 'naive bayes'])) and (not bool(set(algos_3) & set(['k-means clustering', 'apriori', 'linear regression','multilinear regression'])))
True
```
```
>>> 'k-means clustering' in algos_4 and (not bool(set(algos_4) & set(['decision tree Classifier', 'svm', 'logistic regression', 'knn', 'random forest classifier', 'naive bayes', 'multilinear regression', 'apriori', 'linear regression'])))
True
```
```
>>> bool(set(algos_5) & set(['decision tree Classifier', 'svm', 'logistic regression', 'knn', 'random forest classifier', 'naive bayes'])) and (not bool(set(algos_5) & set(['k-means clustering', 'apriori', 'linear regression','multilinear regression'])))
True
```
```
>>> algo_type_2 == 'supervised'
True
```
```
>>> algo_type_4 == 'unsupervised'
True
```
```
>>> task_type_4 == 'clustering'
True
```
```
>>> task_type_5 == 'classification' 
True
```
```
>>> algo_type_5 == 'supervised'
True
```
```
>>> algo_type_3 == 'supervised'
True
```
```
>>> task_type_3 == 'classification' 
True
```
```
>>> task_type_2 == 'classification' 
True
```
```
>>> algo_type_1.lower() == 'supervised'
True
```
```
>>> task_type_1.lower() == 'regression' 
True
```


Tests for q1 (points: 5)
```
>>> 'Sepal Height' in df.columns
True
```
```
>>> 'Sepal Width' in df.columns
True
```
```
>>> 'Petal Height' in df.columns
True
```
```
>>> 'Petal Width' in df.columns
True
```
```
>>> 'Target' in df.columns
True
```
```
>>> df.columns == ['Sepal Height', 'Sepal Width', 'Petal Height', 'Petal Width', 'Target']
array([ True,  True,  True,  True,  True])
```


Tests for q2 (points: 3)
```
>>> set(df['Target'].unique()) == set((0,1,2))
True
```


Tests for q3 (points: 2)
```
>>> df.isna().sum().sum() == 0
True
```


Tests for q4 (points: 5)
```
>>> len(targets) == 3
True
```


Tests for q5 (points: 5)
```
>>> probability(0.1, 0, 1) == 0.3969525474770118
True
```
```
>>> probability(0.1, 1, 1)  == 0.2660852498987548
True
```
```
>>> probability(3.2, 0, 1) == 0.0023840882014648404
True
```
```
>>> probability(-3.2, 0, 1)== 0.0023840882014648404
True
```
```
>>> probability(3.2, 0, 1) == probability(-3.2, 0, 1)
True
```


Tests for q6 (points: 15)
```
>>> result == 0
True
```
```
>>> features = np.array([3.2, 3.0,	1.0,	0.6])
>>> result = naive_bayes(df, features, 'Target')
>>> result == 0
True
```
```
>>> features = np.array([1.2, 1.0,	1.0,	0.6])
>>> result = naive_bayes(df, features, 'Target')
>>> result == 0
True
```
```
>>> features = np.array([10.5, 6.2,	8.9,	7.5])
>>> result = naive_bayes(df, features, 'Target')
>>> result == 2
True
```
```
>>> features = np.array([10, 6,	8,	8])
>>> result = naive_bayes(df, features, 'Target')
>>> result == 2
True
```
```
>>> features = np.array([10.5, 6.2,	8.9,	7.5])
>>> result = naive_bayes(df, features, 'Target')
>>> result == 2
True
```


Tests for q7 (points: 5)
```
>>> x_train.shape == (120, 4)
True
```
```
>>> x_test.shape == (30, 4)
True
```
```
>>> y_train.shape == (120,)
True
```
```
>>> y_test.shape == (30,)
True
```


Tests for q8 (points: 10)
```
>>> accuracy>0.6
True
```
```
>>> accuracy>0.7
True
```
```
>>> accuracy>0.8
True
```
```
>>> accuracy>0.9
True
```
```
>>> for i in range(len(x_test)):
...   res = naive_bayes(train_dataFrame, x_test.iloc[i,:], 'Target')
...   if res!=y_test.iloc[i]:
...     errors = errors + 1
>>> 
>>> (len(x_test) - errors)*100/len(x_test)>0.8
True
```
```
>>> for i in range(len(x_test)):
...   res = naive_bayes(train_dataFrame, x_test.iloc[i,:], 'Target')
...   if res!=y_test.iloc[i]:
...     errors = errors + 1
>>> 
>>> (len(x_test) - errors)*100/len(x_test)>0.9
True
```
```
>>> for i in range(len(x_test)):
...   res = naive_bayes(train_dataFrame, x_test.iloc[i,:], 'Target')
...   if res!=y_test.iloc[i]:
...     errors = errors + 1
>>> 
>>> (len(x_test) - errors)*100/len(x_test)>0.7
True
```
```
>>> for i in range(len(x_test)):
...   res = naive_bayes(train_dataFrame, x_test.iloc[i,:], 'Target')
...   if res!=y_test.iloc[i]:
...     errors = errors + 1
>>> 
>>> (len(x_test) - errors)*100/len(x_test)>0.6
True
```
